```python
# Initialize Otter
import otter
grader = otter.Notebook("cs109a_hw4.ipynb")
```

# <img style="float: left; padding-right: 10px; width: 45px" src="https://raw.githubusercontent.com/Harvard-IACS/2018-CS109A/master/content/styles/iacs.png"> CS109A Introduction to Data Science

## Milestone 3 - Saturn: EDA, Adjusted project plan and Set Goals

**Harvard University**<br/>
**Fall 2023**<br/>

**Watt\'s up with the Weather?** _Correlating City Hall's Electricity Consumption with Weather Patterns_ <br/>

**Group 34:** Robert Escudero, Sean McCabe, Yoona Kim, Lachlain McGranahan, and Calvin Beighle<br/>

**Instructors**: Pavlos Protopapas & Kevin Rader

---

```python
# RUN THIS CELL 
import requests
from IPython.core.display import HTML
styles = requests.get(
    "https://raw.githubusercontent.com/Harvard-IACS/2021-CS109A/master/"
    "themes/static/css/cs109.css"
).text
HTML(styles)
```

Output:
```
<IPython.core.display.HTML object>
```

#### Instructions for Milestone 3
- Summary of the Data: The EDA report should begin with a brief summary of the data. This summary should include information such as the shape of the data, data types, and descriptive statistics such as mean, max, and dtypes. Additionally, you should provide a summary of the features of the data, including histograms, correlation plots, and clustering plots.

- Deeper Understanding of the Data: While basic EDA is important, you should aim to provide a deeper understanding of the data through your analysis. This can be achieved by identifying patterns, trends, class imbalances, and outliers in the data. Additionally, explore the relationships between variables and identify any potential confounding variables that may impact the analysis.

- Meaningful Insights: The EDA report should provide meaningful insights that can be connected back to the problem at hand. These insights should be well-supported by the data and provide actionable recommendations for addressing the problem. You should focus on providing insights that are relevant to the project question and will add value to the final analysis.

- Clean and Labeled Visualization: Visualization is an important component of EDA and should be clean, labeled, and well-presented. You need to ensure that your visualizations are easy to understand and can be included in their final presentation slides or report. Anyone that reads your EDA should be able to understand what is depicted in the plots just by looking at them.

- Data Description: You should provide a detailed description of the data for your project. This should include information about the data source, the data collection process, and any preprocessing steps that were taken. Additionally, you should describe the methods you used to explore the data, including initial explorations, data cleaning, and reconciliation.

- Noteworthy Findings: Summarize the noteworthy findings of their EDA in a clear and concise manner. This can be achieved through the use of visualizations and captions that highlight the most important insights gained through the analysis.

- Project Question: Based on the insights gained through EDA, you should develop a clear project question that will guide your analysis. This question should be well-defined and specific to the problem at hand.

- Baseline Model or Implementation Plan: Finally, you should include a baseline model or a clear plan for its implementation. This can include details on the model architecture, the data used for training and validation, and the evaluation metrics used to assess model performance.

- **You should do a "Restart Kernel and Run All Cells" before submitting to ensure (1) your notebook actually runs and (2) all output is visible**

### Import Libraries

```python
# Import libraries
import seaborn as sns
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.neighbors import KNeighborsRegressor
from sklearn.decomposition import PCA
from sklearn.model_selection import cross_validate
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import r2_score
from sklearn.impute import SimpleImputer, KNNImputer

# pandas tricks for better display
pd.options.display.max_columns = 50  
pd.options.display.max_rows = 500     
pd.options.display.max_colwidth = 100
pd.options.display.precision = 3
```

```python

```

Tests for q1.1 (points: 6)
```
>>> assert type(high_corr) == dict and len(high_corr) == 9,\
... "high_corr should be a dict with 9 entries"
```
```
>>> assert all([isinstance(k, tuple) and len(k) == 2 and isinstance(k[0], str) and isinstance(k[1], str) for k in high_corr.keys()]),\
...     "All keys in high_corr should be tuples of two strings."
```
```
>>> assert all([isinstance(v, (float, np.float_)) and v >= -1 and v <= 1 for v in high_corr.values()]),\
...         "All values in high_corr should be floats between -1 and 1 (inclusive)."
```
```
>>> 
>>> # HIDDEN
>>> def test_high_corr():
...     # Step 1: Use a fixed dataset to compute the expected high_corr dictionary.
...     # (For simplicity, I'm using the same logic as the solution here. In practice, you might use a smaller fixed dataset.)
...     col_names_test = X_train.columns
...     df_corr_test = X_train.corr()
...     itemindex_test = np.where((df_corr_test > 0.99) | (df_corr_test < -0.99))
...     expected_high_corr = {(col_names_test[i],col_names_test[j]): df_corr_test.iloc[i,j] for (i,j) in zip(itemindex_test[0], itemindex_test[1]) if i > j}
...     
...     # Step 2: Compare the student's high_corr dictionary to the expected dictionary.
...     # assert set(high_corr.keys()) == set(expected_high_corr.keys()), "The keys in your high_corr dictionary do not match the expected keys."
...     assert len(high_corr.keys()) == len(expected_high_corr.keys()), "You don't have the same number of keys in the dictionary as you should."
>>> 
>>> test_high_corr()
```


Tests for q1.2 (points: 6)
```
>>> assert linreg_train_r2 > 0.1 and linreg_train_r2 < 0.9,\
... "The true linreg_train_r2 is between 0.1 and 0.9"
```
```
>>> # HIDDEN
>>> assert round(linreg_train_r2, 3) == 0.846 ,\
... "The true linreg_train_r2 is ~0.846"
```
```
>>> assert linreg_test_r2 > 0.1 and linreg_test_r2 < 0.9,\
... "The true linreg_test_r2 is between 0.1 and 0.9"
```
```
>>> assert round(linreg_test_r2, 3) == 0.219 ,\
... "The true linreg_test_r2 is ~0.219"
```


Tests for q1.3.1 (points: 14)
```
>>> assert best_k > 4 and best_k < 15,\
... "The true best_k is between 4 and 15"
```
```
>>> # HIDDEN
>>> assert best_k == 8 ,\
... "The true best_k is 8"
```


Tests for q1.3.2 (points: 6)
```
>>> assert pcr_train_r2 > 0.2 and pcr_train_r2 < 0.8,\
... "The true pcr_train_r2 is between 0.2 and 0.8"
```
```
>>> # HIDDEN
>>> assert round(pcr_train_r2, 3) == 0.638 ,\
... "The true pcr_train_r2 is ~0.638"
```
```
>>> assert pcr_test_r2 > 0.2 and pcr_test_r2 < 0.8,\
... "The true pcr_test_r2 is between 0.2 and 0.8"
```
```
>>> # HIDDEN
>>> assert round(pcr_test_r2, 3) == 0.686 ,\
... "The true pcr_test_r2 is ~0.686"
```


Tests for q2.1 (points: 2)
```
>>> assert X_transformed.shape == (38, 2),\
... "The transformed data should have shape (38, 2)."
```
```
>>> assert np.isclose(X_transformed.mean(), 0),\
... "The mean of the transformed data should be 0 (or very close given machine precision)."
```
```
>>> # HIDDEN
>>> def test_transformed():
...     df = pd.read_csv("data/OECD_well-being.csv", index_col = 0)
...     country, X = df['Country'], df.drop(columns='Country').values
...     scaler = StandardScaler()
...     X_std= scaler.fit_transform(X)
...     pca = PCA(n_components=2)
...     X_transformed_sol = pca.fit_transform(X_std)
...     assert np.allclose(X_transformed, X_transformed_sol),\
...     "Transformed data is incorrect. It doesn't match the results in the solutions."
```


Tests for q3.1 (points: 4)
```
>>> def test_n_rows_with_missingness():
...     assert isinstance(n_rows_with_missingness, (int, np.integer)), "n_rows_with_missingness should be an integer" 
...     assert n_rows_with_missingness >=0, "n_rows_with_missingness should be non-negative" 
>>> test_n_rows_with_missingness()
```
```
>>> # HIDDEN
>>> assert n_rows_with_missingness == X_train.isna().max(axis=1).sum(),\
... "n_rows_with_missingness is incorrect"
```
```
>>> def test_n_cols_with_missingness():
...     assert isinstance(n_cols_with_missingness, (int, np.integer)), "n_cols_with_missingness should be an integer" 
...     assert n_cols_with_missingness >=0, "n_cols_with_missingness should be non-negative" 
>>> test_n_cols_with_missingness()
```
```
>>> # HIDDEN
>>> assert n_cols_with_missingness == X_train.isna().max(axis=0).sum(),\
... "n_cols_with_missingness is incorrect"
```
```
>>> def test_col_missingness():
...     col_missingness_sol = X_train.isna().sum(axis=0)
...     assert type(col_missingness) == pd.Series, "col_missingness should be a Pandas Series object"
...     assert all([type(i) == str for i in col_missingness.index]), "Indices of col_missingness should be strings (i.e., the column names)"
...     assert set(col_missingness.index) == set(X_train.columns), "Indices of col_missingness should be the names of the X_train columns."
...     assert all([np.issubdtype(v, np.integer) for v in col_missingness.values]), "Values in col_missingness should be some type of integer"
>>> test_col_missingness()
```
```
>>> # HIDDEN
>>> col_missingness_sol = X_train.isna().sum(axis=0)
>>> assert pd.Series.equals(col_missingness[X_train.columns], col_missingness_sol),\
... "You have the wrong numbers of missing values for the entries in your Series"
```


Tests for q4.1 (points: 12)
```
>>> assert linreg_meanimp_r2 > 0.1 and linreg_meanimp_r2 < 0.2,\
... "The true linreg_meanimp_r2 is between 0.1 and 0.2"
```
```
>>> # HIDDEN
>>> assert round(linreg_meanimp_r2, 3) == 0.189 ,\
... "The true linreg_meanimp_r2 is ~0.189"
```
```
>>> assert knn_meanimp_r2 > 0.35 and knn_meanimp_r2 < 0.45,\
... "The true knn_meanimp_r2 is between 0.35 and 0.45"
```
```
>>> # HIDDEN
>>> assert round(knn_meanimp_r2, 3) == 0.401 ,\
... "The true knn_meanimp_r2 is between 0.401"
```


Tests for q4.2 (points: 12)
```
>>> assert linreg_knnimp_r2 > 0.2 and linreg_knnimp_r2 < 0.3,\
... "The true linreg_knnimp_r2 is between 0.2 and 0.3"
```
```
>>> # HIDDEN
>>> assert round(linreg_knnimp_r2, 3) == 0.205 ,\
... "The true linreg_knnimp_r2 is ~0.205"
```
```
>>> assert knn_knnimp_r2 > 0.55 and knn_knnimp_r2 < 0.65,\
... "The true knn_knnimp_r2 is between 0.55 and 0.65"
```
```
>>> # HIDDEN
>>> assert round(knn_knnimp_r2, 3) == 0.622 ,\
... "The true knn_knnimp_r2 is between 0.622"
```


Tests for q4.3 (points: 12)
```
>>> assert linreg_indic_r2 > 0.2 and linreg_indic_r2 < 0.5,\
... "The true linreg_indic_r2 is between 0.2 and 0.5"
```
```
>>> # HIDDEN
>>> assert round(linreg_indic_r2, 3) == 0.395 ,\
... "The true linreg_knnimp_r2 is ~0.395"
```
```
>>> assert knn_indic_r2 > 0.55 and knn_indic_r2 < 0.75,\
... "The true knn_indic_r2 is between 0.55 and 0.75"
```
```
>>> # HIDDEN
>>> assert round(knn_indic_r2, 3) == 0.653 ,\
... "The true knn_indic_r2 is between ~0.653"
```
